# Aufgabe 1
Vervollst채ndie die folgende Funktion, welche die Stunde aus einer erhaltenen E-Mail herausfiltert und die H채ufigkeit mithilfe von yield ausgibt.

Datei ist in folgendem Format geschrieben:

~~~~
From max.mustermann@fh-swf.de Sat Jan 5 09:14:16 2019
Message
Hallo,

das ist meine super E-Mail.
From susanne.mustermann@fh-swf.de Sun Jan 6 16:31:02 2019
Message
Danke f체r die E-Mail!
~~~~

Beispiel Ausgabe:<br>
('09', 1)<br>
('16', 1)<br>
<br>
Die Stunde steht immer in der Zeile, welche mit 'From ' beginnt.

<b>Hinweis</b>:
- die Zeiten m체ssen sortiert nach der Stunde ausgegeben werden
- die Ausgabe erfolgt als Tuple im Format (stunde, stunde_vorkommen)
- es kann vorkommen, dass in einer 'From ' Zeile bis zu 2 Leerzeichen hintereinander vorkommen

<!-- grade_id: cell-d5f0431c8aca09d5, points: 0, grade: false, solution: true -->
```python
import collections

def get_hour_frequency(datei: str) -> tuple:
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
for i in get_hour_frequency("email_test.txt"):
    print(i)
```

Output:
```
File "<ipython-input-39-f9551a50a9b2>", line 8
    for i in get_hour_frequency("email_test.txt"):
      ^
IndentationError: expected an indented block
```

<!-- grade_id: cell-aacd27206f9a7ad1, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

answer_test = [('09', 1), ('16', 1)]
answer = [('00', 408), ('01', 392), ('02', 431), ('03', 439), ('04', 432), ('05', 418), ('06', 436), ('07', 389), ('08', 425), ('09', 408), ('10', 414), ('11', 426), ('12', 422), ('13', 401), ('14', 427), ('15', 446), ('16', 394), ('17', 421), ('18', 428), ('19', 428), ('20', 394), ('21', 428), ('22', 381), ('23', 412)]

for i, k in enumerate(get_hour_frequency("email_test.txt")):
    assert_equal(k, answer_test[i])
    
for i, k in enumerate(get_hour_frequency("emails.txt")):
    assert_equal(k, answer[i])
```

Output:
```
[["09", 1], ["16", 1]]
NameError                                 Traceback (most recent call last)
<ipython-input-1-d3476c1ac836> in <module>
     14 answer = [('00', 408), ('01', 392), ('02', 431), ('03', 439), ('04', 432), ('05', 418), ('06', 436), ('07', 389), ('08', 425), ('09', 408), ('10', 414), ('11', 426), ('12', 422), ('13', 401), ('14', 427), ('15', 446), ('16', 394), ('17', 421), ('18', 428), ('19', 428), ('20', 394), ('21', 428), ('22', 381), ('23', 412)]
     15 
---> 16 for i, k in enumerate(get_hour_frequency("email_test.txt")):
     17     assert_equal(k, answer_test[i])
     18 

NameError: name 'get_hour_frequency' is not defined
```