# Aufgabe 3

Vervollständie die folgende Funktion, um die Häufigkeiten eines Buchstaben zu ermitteln. Es sollen nur Buchstaben von a-z gewertet werden. Jeder Buchstabe muss zum Kleinbuchstaben konvertiert werden.

Beispiel:

~~~~
Hallo!
~~~~

Beispiel Ausgabe als Dictionary:

~~~~
frequency = {
    'a': 1,
    'h': 1,
    'l': 2,
    'o': 1
}
~~~~

<b>Info:</b> Stimmt die Verteilung der Buchstabenhäufigkeiten in einem Längeren Text mit der Tabelle überein: https://wikipedia.org/wiki/Letter_frequencies?

<!-- grade_id: cell-f54ff9e94c457e7e, points: 0, grade: false, solution: true -->
```python
import collections

def get_letter_frequency(eingabe: str) -> dict:
    frequency = dict()
    
    ### BEGIN SOLUTON
    ### END SOLUTION
    
    return frequency

print(get_letter_frequency('Hallo!'))
print(get_letter_frequency('Am Flughafen Familie Müller plant ihren Urlaub. Sie geht in ein Reisebüro und lässt sich von einem Angestellten beraten. Als Reiseziel wählt sie Mallorca aus. Familie Müller bucht einen Flug auf die Mittelmeerinsel. Sie bucht außerdem zwei Zimmer in einem großen Hotel direkt am Strand. Familie Müller badet gerne im Meer. Am Abflugtag fahren Herr und Frau Müller mit ihren beiden Kindern im Taxi zum Flughafen. Dort warten schon viele Urlauber. Alle wollen nach Mallorca fliegen. Familie Müller hat viel Gepäck dabei: drei große Koffer und zwei Taschen. Die Taschen sind Handgepäck. Familie Müller nimmt sie mit in das Flugzeug. Am Flugschalter checkt die Familie ein und erhält ihre Bordkarten. Die Angestellte am Flugschalter erklärt Herrn Müller den Weg zum Flugsteig. Es ist nicht mehr viel Zeit bis zum Abflug. Familie Müller geht durch die Sicherheitskontrolle. Als alle das richtige Gate erreichen, setzen sie sich in den Wartebereich. Kurz darauf wird ihre Flugnummer aufgerufen und Familie Müller steigt mit vielen anderen Passagieren in das Flugzeug nach Mallorca. Beim Starten fühlt sich Herr Müller nicht wohl. Ihm wird ein wenig übel. Nach zwei Stunden landet das Flugzeug. Am Gepäckband warten alle Passagiere noch auf ihr fehlendes Gepäck. Danach kann endlich der Urlaub beginnen.'))
```

Output:
```
{}
{}
```

<!-- grade_id: cell-819ba6b7a8e37178, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(get_letter_frequency("Hallo!"), {'a': 1, 'h': 1, 'l': 2, 'o': 1})
assert_equal(get_letter_frequency('"Der Charakter ruht auf der Persönlichkeit, nicht auf den Talenten." ~ Johann Wolfgang von Goethe'), {'a': 7, 'c': 3, 'd': 3, 'e': 10, 'f': 3, 'g': 3, 'h': 6, 'i': 3, 'j': 1, 'k': 2, 'l': 3, 'n': 9, 'o': 4, 'p': 1, 'r': 6, 's': 1, 't': 7, 'u': 3, 'v': 1, 'w': 1})
assert_equal(get_letter_frequency(" "), {})
assert_equal(get_letter_frequency('Am Flughafen Familie Müller plant ihren Urlaub. Sie geht in ein Reisebüro und lässt sich von einem Angestellten beraten. Als Reiseziel wählt sie Mallorca aus. Familie Müller bucht einen Flug auf die Mittelmeerinsel. Sie bucht außerdem zwei Zimmer in einem großen Hotel direkt am Strand. Familie Müller badet gerne im Meer. Am Abflugtag fahren Herr und Frau Müller mit ihren beiden Kindern im Taxi zum Flughafen. Dort warten schon viele Urlauber. Alle wollen nach Mallorca fliegen. Familie Müller hat viel Gepäck dabei: drei große Koffer und zwei Taschen. Die Taschen sind Handgepäck. Familie Müller nimmt sie mit in das Flugzeug. Am Flugschalter checkt die Familie ein und erhält ihre Bordkarten. Die Angestellte am Flugschalter erklärt Herrn Müller den Weg zum Flugsteig. Es ist nicht mehr viel Zeit bis zum Abflug. Familie Müller geht durch die Sicherheitskontrolle. Als alle das richtige Gate erreichen, setzen sie sich in den Wartebereich. Kurz darauf wird ihre Flugnummer aufgerufen und Familie Müller steigt mit vielen anderen Passagieren in das Flugzeug nach Mallorca. Beim Starten fühlt sich Herr Müller nicht wohl. Ihm wird ein wenig übel. Nach zwei Stunden landet das Flugzeug. Am Gepäckband warten alle Passagiere noch auf ihr fehlendes Gepäck. Danach kann endlich der Urlaub beginnen.'), {'a': 79, 'b': 19, 'c': 32, 'd': 40, 'e': 160, 'f': 34, 'g': 38, 'h': 48, 'i': 94, 'k': 13, 'l': 87, 'm': 50, 'n': 78, 'o': 17, 'p': 7, 'r': 76, 's': 44, 't': 54, 'u': 43, 'v': 5, 'w': 13, 'x': 1, 'z': 14})
```

Output:
```
/opt/tljh/user/lib/python3.7/unittest/case.py in assertDictEqual(self, d1, d2, msg)
   1136                            pprint.pformat(d2).splitlines())))
   1137             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1138             self.fail(self._formatMessage(msg, standardMsg))
   1139 
   1140     def assertDictContainsSubset(self, subset, dictionary, msg=None):

/opt/tljh/user/lib/python3.7/unittest/case.py in fail(self, msg)
    678     def fail(self, msg=None):
    679         """Fail immediately, with the given message."""
--> 680         raise self.failureException(msg)
    681 
    682     def assertFalse(self, expr, msg=None):

AssertionError: {} != {'a': 1, 'h': 1, 'l': 2, 'o': 1}
- {}
+ {'a': 1, 'h': 1, 'l': 2, 'o': 1}
```