# Zinseszins


### Formel
$K_{n}=K*(1+\frac{p}{100})^n$

Hierbei ist:<br>
>$K_n$ das Kapital nach der Verzinsung<br>
$K$ das Startkapital<br>
$p$ die Zinszahl<br>
$n$ die Anzahl der Jahre<br>

### Quelle
Eine Erklärung des Zinseszins und die Herleitung der Formel sind auf
https://www.gut-erklaert.de/mathematik/zinseszins-berechnen-formel-erklaerung.html
zu finden

## Aufgabe
Schreibe eine Funktion, um den Zinseszins zu berechnen.

### Hinweis
Achte darauf, dass du nicht in der Zeit zurückreisen kannst und dein Kapital im positiven startet. <br>
Wirf eine <b>Exception</b> im Falle eines Fehlers.

<!-- grade_id: cell-5dd5b6bea6f8e411, points: 0, grade: false, solution: true -->
```python
import math

def calculate_compound_interest(K: float, p: float, n: int):
    """
    Funktion um den Zinseszins zu berechnen
    
    :returns: Kapital nach der Verzinsung
    :param K: Startkapital
    :param p: Zinszahl
    :param n: Anzahl der Jahre
    """
    
    K_n: float = None
    
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    
    return K_n
```

<!-- grade_id: cell-5a36737a9698db49, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_raises

assert_equal(calculate_compound_interest(0, 10, 0), 0.0)

assert_raises(Exception, calculate_compound_interest, -1, 10, 0)
assert_raises(Exception, calculate_compound_interest, -1, 10, 1)

assert_equal(calculate_compound_interest(1, -1, 0), 1.0)
assert_equal(calculate_compound_interest(1, -1, 1), 0.99)

assert_equal(calculate_compound_interest(2000, 3, 0), 2000.0)
assert_equal(calculate_compound_interest(2000, 3, 4), 2251.02)
assert_equal(calculate_compound_interest(2000, 3, 5), 2318.55)
assert_equal(calculate_compound_interest(2000, 3, 6), 2388.10)

### BEGIN HIDDEN TESTS
assert_raises(Exception, calculate_compound_interest, 1, -1, -1)
assert_raises(Exception, calculate_compound_interest, -1, 10, -1)
assert_equal(calculate_compound_interest(0, 10, 1), 0.0)
assert_raises(Exception, calculate_compound_interest, 0, 10, -1)
assert_equal(calculate_compound_interest(2000, 9, 100), 11058081.58)
assert_equal(calculate_compound_interest(10000, -1, 10), 9043.82)
assert_equal(calculate_compound_interest(20000, 3, 1), 20600.0)
### END HIDDEN TESTS
```

## Aufgabe
Runde das Ergebnis auf 2 Nachkommastellen

<!-- grade_id: cell-deee56cf4be14bc6, points: 0, grade: false, solution: true -->
```python
def calculate_compound_interest_rounded(K: float, p: float, n: int):
    """
    Funktion um den Zinseszins zu berechnen
    
    :returns: Kapital nach der Verzinsung gerundet auf 2 Nachkommastellen
    :param K: Startkapital
    :param p: Zinszahl
    :param n: Anzahl der Jahre
    """
    
    K_n: float = None
    
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    
    return K_n
```

<!-- grade_id: cell-1f0b9a11b788fb5b, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_raises

assert_equal(calculate_compound_interest_rounded(0, 10, 0), 0.0)

assert_raises(Exception, calculate_compound_interest_rounded, -1, 10, 0)
assert_raises(Exception, calculate_compound_interest_rounded, -1, 10, 1)

assert_equal(calculate_compound_interest_rounded(1, -1, 0), 1.0)
assert_equal(calculate_compound_interest_rounded(1, -1, 1), 0.99)

assert_equal(calculate_compound_interest_rounded(2000, 3, 0), 2000.0)
assert_equal(calculate_compound_interest_rounded(2000, 3, 4), 2251.02)
assert_equal(calculate_compound_interest_rounded(2000, 3, 5), 2318.55)
assert_equal(calculate_compound_interest_rounded(2000, 3, 6), 2388.10)

### BEGIN HIDDEN TESTS
assert_raises(Exception, calculate_compound_interest_rounded, 1, -1, -1)
assert_raises(Exception, calculate_compound_interest_rounded, -1, 10, -1)
assert_equal(calculate_compound_interest_rounded(0, 10, 1), 0.0)
assert_raises(Exception, calculate_compound_interest_rounded, 0, 10, -1)
assert_equal(calculate_compound_interest_rounded(2000, 9, 100), 11058081.58)
assert_equal(calculate_compound_interest_rounded(10000, -1, 10), 9043.82)
assert_equal(calculate_compound_interest_rounded(20000, 3, 1), 20600.0)
### END HIDDEN TESTS
```

```python

```