# Aufgabe 4

Vervollständige die folgenden Funktionen, welche das Minimum und das Maximum ausgeben. Gib <b>None</b> zurück, wenn kein Minimum oder Maximum gefunden werden kann.

<!-- grade_id: cell-f285d88888b6d2d9, points: 0, grade: false, solution: true -->
```python
from typing import List

def minimum(liste: List[int]) -> int:
    ### BEGIN SOLUTION
    ### END SOLUTION
    
def maximum(liste: List[int]) -> int:
    ### BEGIN SOLUTION
    ### END SOLUTION
    
print(minimum([4, 10, 3]))
print(maximum([4, 10, 3]))
```

Output:
```
File "<ipython-input-5-db88c27746e1>", line 7
    def maximum(liste: List[int]) -> int:
      ^
IndentationError: expected an indented block
```

<!-- grade_id: cell-53441d2a2d9261a3, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(minimum([10, 340, 50, 20, 10, 30, 50, 11, 2, 56, 13]), 2)
assert_equal(minimum([40, 12, 12, 0, 99, -1, 4]), -1)
assert_equal(minimum([-13, -314, 0, -53, -98, -1]), -314)
assert_equal(minimum([1, 2, 3, 4, 5]), 1)

assert_equal(maximum([19837, 9182, 1927, 784, 9182, 1, 0]), 19837)
assert_equal(maximum([982, 1289, 1982, 491, 19]), 1982)
assert_equal(maximum([0, 1, 2, 3, 4, 5]), 5)
assert_equal(maximum([-1, -13, -52, -1, -985]), -1)
assert_equal(maximum([0, 1]), 1)
assert_equal(maximum([]), None)
assert_equal(minimum([]), None)
```