# Aufgabe 1

Verfollständige die folgende Funktion, welche das Monatsgehalt eines Mitarbeiters zurückgibt. Arbeitsstunden über 40 Stunden werden mit dem Faktor 1.5 multipliziert.<br>
Runde das Ergebnis aus 2 Nachkommastellen!

<!-- grade_id: cell-d61869b6e5e4ec5b, points: 0, grade: false, solution: true -->
```python
def calculate_salary(arbeitsstunden: float, stundenlohn: float) -> float:
    ### BEGIN SOLUTION
    ### END SOLUTION
    
print(calculate_salary(40, 1.2))
```

Output:
```
File "<ipython-input-15-2a74d391fdfb>", line 5
    print(calculate_salary(40, 1.2))
        ^
IndentationError: expected an indented block
```

<!-- grade_id: cell-7adafe3454ea1a76, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_raises

assert_equal(calculate_salary(45, 10), 475)
assert_equal(calculate_salary(0, 0), 0)
assert_equal(calculate_salary(0, 10.4), 0)
assert_equal(calculate_salary(200, 12.6), 3528)
assert_equal(calculate_salary(201, 12.6), 3546.9)
assert_equal(calculate_salary(39, 22.1), 861.9)
assert_equal(calculate_salary(40, 1.2), 48.0)

assert_raises(ValueError, calculate_salary, -8.5, 12.5)
assert_raises(ValueError, calculate_salary, 40, -10.62)
assert_raises(ValueError, calculate_salary, -100, -12.5)
assert_raises(ValueError, calculate_salary, -1.1, -0.1)
```

```python

```