# Aufgabe 1

Übung zum aufwärmen. Schreibe folgenden Programmcode korrigiert in die untere Codezelle:

~~~~
func value_greater_five(int value) {
    int my_int = 5
    
    if (value > my_int) {
        return "greater"
    }
    
    return "lower"
}
~~~~

<!-- grade_id: cell-52d4b9650440974a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION

print(value_greater_five(1))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-6-12f315118441> in <module>
      2 ### END SOLUTION
      3 
----> 4 print(value_greater_five(1))

NameError: name 'value_greater_five' is not defined
```

<!-- grade_id: cell-49073f0643ea7b5d, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(value_greater_five(-1), "lower")
assert_equal(value_greater_five(0), "lower")
assert_equal(value_greater_five(2), "lower")
assert_equal(value_greater_five(5), "lower")
assert_equal(value_greater_five(10), "greater")
assert_equal(value_greater_five(100), "greater")
assert_equal(value_greater_five(1000), "greater")
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-4-418353f77df2> in <module>
      1 from nose.tools import assert_equal
      2 
----> 3 assert_equal(value_greater_five(-1), "lower")
      4 assert_equal(value_greater_five(0), "lower")
      5 assert_equal(value_greater_five(2), "lower")

NameError: name 'value_greater_five' is not defined
```

```python

```