# Aufgabe 4

Vervollständige die folgende Funktionen, um Temperaturen zwischen Celsius und Fahrenheit zu konvertiert.<br>
Runde das Ergebnis auf 2 Nachkommastellen!

<!-- grade_id: cell-128e70779f2362a7, points: 0, grade: false, solution: true -->
```python
def fahrenheit_to_celsius(fahrenheit: float) -> float:
    ### BEGIN SOLUTION
    ### END SOLUTION
    
def celsius_to_fahrenheit(celsius: float) -> float:
    ### BEGIN SOLUTION
    ### END SOLUTION
    
print(fahrenheit_to_celsius(102))
print(celsius_to_fahrenheit(30.3))
```

<!-- grade_id: cell-a03c0cfdac04ebdc, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(fahrenheit_to_celsius(98.6), 37)
assert_equal(fahrenheit_to_celsius(0), -17.78)
assert_equal(fahrenheit_to_celsius(-100), -73.33)
assert_equal(fahrenheit_to_celsius(50), 10)
assert_equal(fahrenheit_to_celsius(200.41), 93.56)

assert_equal(celsius_to_fahrenheit(0), 32)
assert_equal(celsius_to_fahrenheit(37), 98.6)
assert_equal(celsius_to_fahrenheit(15.6), 60.08)
assert_equal(celsius_to_fahrenheit(100), 212)
assert_equal(celsius_to_fahrenheit(-10), 14)
assert_equal(celsius_to_fahrenheit(-105.8), -158.44)
```