# Aufgabe 3

Vervollständige die folgende Funktion, um den Bruttolohn zu berechnen. Gib einen <b>ValueError</b> aus, wenn ein eingegebener Wert nicht passt.

<!-- grade_id: cell-f21943c9a9dd0be9, points: 0, grade: false, solution: true -->
```python
def calculate_gross(arbeitsstunden: float, stundensatz: float) -> float:
    ### BEGIN SOLUTION
    ### END SOLUTION
    
print(calculate_gross(10, 10))
```

Output:
```
File "<ipython-input-2-8500beef3772>", line 5
    print(calculate_gross(10, 10))
        ^
IndentationError: expected an indented block
```

<!-- grade_id: cell-f95d52d817168931, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_raises

assert_equal(calculate_gross(35, 2.75), 96.25)
assert_equal(calculate_gross(40, 11.90), 476)
assert_equal(calculate_gross(200, 22.45), 4490)
assert_equal(calculate_gross(201, 22.01), 4424.01)
assert_equal(calculate_gross(0, 12.66), 0)
assert_raises(ValueError, calculate_gross, -1, 12.66)
assert_raises(ValueError, calculate_gross, 0, -12.66)
assert_raises(ValueError, calculate_gross, -10, -12.66)
```