# Aufgabe 1

Vervollständie die folgende Funktion, welche die Wörter aus einer Datei in ein Dictionary speichert, welches als Wert die Anzahl des Wortes in der Datei angibt.

Hinweis:
- folgende Zeichen müssen entfernt werden: "<b>(),.{}?\\:</b>"
- der Text muss getrimmt werden
- gültige Wörter haben eine Länge von mindestens 2 Buchstaben
- alle Buchstaben der Wörter müssen zu Kleinbuchstaben geändert werden

<!-- grade_id: cell-391a82abbbf16ec2, points: 0, grade: false, solution: true -->
```python
def count_words(datei: str) -> dict:
    words = dict()
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
    return words

print(count_words("words.txt"))
```

Output:
```
{}
```

<!-- grade_id: cell-b8cad191f79267db, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal
import json

def myHash(text: str):
    hash=0
    for ch in text:
        hash = ( hash*281  ^ ord(ch)*997) & 0xFFFFFFFF
    
    return hex(hash)[2:].upper().zfill(8)

assert_equal(myHash(json.dumps(count_words("words.txt"))), 'A2D8D6B1')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-19f27249d518> in <module>
      9     return hex(hash)[2:].upper().zfill(8)
     10 
---> 11 assert_equal(myHash(json.dumps(count_words("words.txt"))), 'A2D8D6B1')

NameError: name 'count_words' is not defined
```