# Aufgabe 2

Vervollst채ndige die folgende Funktion, welche aus E-Mails den Wochentag herausfiltert und die Anzahl der jeweiligen Wochentage als Dictionary ausgibt.
Die Datei ist in folgendem Format geschrieben:
~~~~
From max.mustermann@fh-swf.de Sat Jan 5 09:14:16 2019
Message
Hallo,

das ist meine super E-Mail.
From susanne.mustermann@fh-swf.de Sun Jan 6 16:31:02 2019
Message
Danke f체r die E-Mail!
~~~~

Beispiel Ausgabe: <b>{'Mon': 0, 'Tue': 0, 'Wen': 0, 'Thu': 0, 'Fri': 0, 'Sat': 1, 'Sun': 1}</b>
<br>Der Wochentag steht immer in der Zeile, welche mit <b>'From '</b> beginnt.

<!-- grade_id: cell-84563f200384f37e, points: 0, grade: false, solution: true -->
```python
def count_weekdays(datei: str) -> dict:
    weekday_count = {
        'Mon': 0,
        'Tue': 0,
        'Wed': 0,
        'Thu': 0,
        'Fri': 0,
        'Sat': 0,
        'Sun': 0
    }
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
    return weekday_count

print(count_weekdays("email_test.txt"))
```

Output:
```
{'Mon': 0, 'Tue': 0, 'Wed': 0, 'Thu': 0, 'Fri': 0, 'Sat': 0, 'Sun': 0}
```

<!-- grade_id: cell-9b9067d2bce228e3, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

answer_test = {'Mon': 0, 'Tue': 0, 'Wed': 0, 'Thu': 0, 'Fri': 0, 'Sat': 1, 'Sun': 1}
answer_mails = {'Mon': 1423, 'Tue': 1515, 'Wed': 1395, 'Thu': 1429, 'Fri': 1425, 'Sat': 1399, 'Sun': 1414}

assert_equal(count_weekdays("email_test.txt"), answer_test)
assert_equal(count_weekdays("emails.txt"), answer_mails)
```

Output:
```
/opt/tljh/user/lib/python3.7/unittest/case.py in assertDictEqual(self, d1, d2, msg)
   1136                            pprint.pformat(d2).splitlines())))
   1137             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1138             self.fail(self._formatMessage(msg, standardMsg))
   1139 
   1140     def assertDictContainsSubset(self, subset, dictionary, msg=None):

/opt/tljh/user/lib/python3.7/unittest/case.py in fail(self, msg)
    678     def fail(self, msg=None):
    679         """Fail immediately, with the given message."""
--> 680         raise self.failureException(msg)
    681 
    682     def assertFalse(self, expr, msg=None):

AssertionError: {'Mon': 0, 'Tue': 0, 'Wed': 0, 'Thu': 0, 'Fri': 0, 'Sat': 0, 'Sun': 0} != {'Mon': 0, 'Tue': 0, 'Wed': 0, 'Thu': 0, 'Fri': 0, 'Sat': 1, 'Sun': 1}
- {'Fri': 0, 'Mon': 0, 'Sat': 0, 'Sun': 0, 'Thu': 0, 'Tue': 0, 'Wed': 0}
?                             ^         ^

+ {'Fri': 0, 'Mon': 0, 'Sat': 1, 'Sun': 1, 'Thu': 0, 'Tue': 
```

# Aufgabe 3

Vervollst채ndie die folgende Funktion, welche den Wochentag mit den meisten Vorkommen aufweist.

Die Datei ist in folgendem Format geschrieben:

~~~~
From max.mustermann@fh-swf.de Sat Jan 5 09:14:16 2019
Message
Hallo,

das ist meine super E-Mail.
From susanne.mustermann@fh-swf.de Sun Jan 6 16:31:02 2019
Message
Danke f체r die E-Mail!
~~~~

Beispiel Ausgabe: 'Sun'

Im Fall mehrerer gleicher Aufkommen, wird der letzte Wochentag mit dieser Anzahl ausgegeben.
<br>Der Wochentag steht immer in der Zeile, welche mit <b>'From '</b> beginnt.

<!-- grade_id: cell-ee26b4f8bfc9932b, points: 0, grade: false, solution: true -->
```python
def max_weekday(datei: str) -> str:
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
    
print(max_weekday("email_test.txt"))
```

<!-- grade_id: cell-37598bf0eb9f9717, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(max_weekday("email_test.txt"), 'Sun')
assert_equal(max_weekday("emails.txt"), 'Tue')
```