# Aufgabe 1

Vervollständie die folgende Funktion, welche die Wörter aus einer Datei in ein Dictionary speichert, welches als Wert die Anzahl des Wortes in der Datei angibt.

Hinweis:
- folgende Zeichen müssen entfernt werden: "<b>(),.{}?\\:</b>"
- der Text muss getrimmt werden
- gültige Wörter haben eine Länge von mindestens 2 Buchstaben
- alle Buchstaben der Wörter müssen zu Kleinbuchstaben geändert werden

<!-- grade_id: cell-391a82abbbf16ec2, points: 0, grade: false, solution: true -->
```python
def count_words(datei: str) -> dict:
    words = dict()
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
    return words

print(count_words("words.txt"))
```

Output:
```
{}
```

<!-- grade_id: cell-b8cad191f79267db, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

answer = {'das': 1, 'schreiben': 2, 'von': 4, 'programmen': 1, 'oder': 2, 'programmieren': 1, 'ist': 2, 'eine': 2, 'sehr': 2, 'kreative': 1, 'und': 7, 'lohnende': 1, 'tätigkeit': 1, 'wir': 8, 'können': 3, 'programme': 1, 'aus': 2, 'vielen': 2, 'gründen': 1, 'angefangen': 1, 'mit': 2, 'dem': 3, 'ziel': 1, 'damit': 1, 'den': 2, 'lebensunterhalt': 1, 'zu': 6, 'verdienen': 1, 'ein': 2, 'schwieriges': 1, 'datenanalyseproblem': 1, 'lösen': 1, 'spaß': 1, 'haben': 3, 'um': 2, 'jemand': 1, 'anderem': 1, 'bei': 2, 'der': 3, 'lösung': 1, 'eines': 1, 'problems': 1, 'helfen': 1, 'dieses': 1, 'buch': 1, 'geht': 1, 'davon': 1, 'dass': 4, 'em': 1, 'jeder': 1, 'wissen': 1, 'sollte': 1, 'wie': 1, 'man': 4, 'programmiert': 1, 'sobald': 1, 'die': 10, 'programmierung': 1, 'beherrscht': 1, 'herausfindet': 1, 'was': 3, 'neugewonnenen': 1, 'fähigkeiten': 1, 'machen': 1, 'kann': 2, 'sind': 3, 'in': 3, 'unserem': 2, 'täglichen': 1, 'leben': 1, 'computern': 2, 'umgeben': 1, 'laptops': 1, 'bis': 1, 'hin': 1, 'handys': 1, 'nehmen': 1, 'diese': 2, 'computer': 5, 'als': 3, 'unsere': 2, 'persönlichen': 2, 'assistenten': 2, 'war': 1, 'viele': 1, 'dinge': 3, 'für': 1, 'uns': 4, 'erledigen': 2, 'hardware': 2, 'unseren': 1, 'heutigen': 1, 'im': 1, 'wesentlichen': 1, 'so': 1, 'gebaut': 1, 'sie': 1, 'ständig': 1, 'frage': 1, 'stellt': 1, 'soll': 3, 'ich': 1, 'nächstes': 2, 'tun': 2, 'programmierer': 1, 'fügen': 1, 'betriebssystem': 1, 'reihe': 1, 'anwendungen': 1, 'hinzu': 1, 'schon': 1, 'einen': 1, 'digitalen': 1, 'problemen': 1, 'des': 1, 'alltags': 1, 'nützlich': 1, 'sein': 2, 'schnell': 1, 'riesige': 1, 'mengen': 1, 'an': 1, 'speicher': 1, 'könnten': 2, 'hilfreich': 1, 'wenn': 2, 'nur': 1, 'sprache': 2, 'beherrschen': 1, 'würden': 1, 'erklären': 1, 'er': 2, 'kennen': 1, 'sagen': 1, 'namen': 1, 'aufgaben': 1, 'sich': 1, 'wiederholen': 1, 'interessanterweise': 1, 'am': 1, 'besten': 1, 'oft': 1, 'menschen': 1, 'langweilig': 1, 'stumpfsinnig': 1, 'finden': 1}
assert_equal(count_words("words.txt"), answer)
```

Output:
```
/opt/tljh/user/lib/python3.7/unittest/case.py in assertDictEqual(self, d1, d2, msg)
   1136                            pprint.pformat(d2).splitlines())))
   1137             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1138             self.fail(self._formatMessage(msg, standardMsg))
   1139 
   1140     def assertDictContainsSubset(self, subset, dictionary, msg=None):

/opt/tljh/user/lib/python3.7/unittest/case.py in fail(self, msg)
    678     def fail(self, msg=None):
    679         """Fail immediately, with the given message."""
--> 680         raise self.failureException(msg)
    681 
    682     def assertFalse(self, expr, msg=None):

AssertionError: {} != {'das': 1, 'schreiben': 2, 'von': 4, 'prog[1899 chars]': 1}
Diff is 2389 characters long. Set self.maxDiff to None to see it.
```