# Aufgabe 2

Vervollständie die folgende Funktion, welche das Inhaltsverzeichnis des Grundgesetzes von der Seite des Deutschen Bundestages auslesen soll. Dazu soll auf die Seite https://www.bundestag.de/parlament/aufgaben/rechtsgrundlagen/grundgesetz zugegriffen werden und die unten stehende Tabelle in folgendes Format umgewandelt werden:<br>

my_dictionary = {<br>
    "Seitennummern" : "Titel",<br>
    ...<br>
}

<b>Info</b>: Es kann sich eine Bibliothek ausgesucht werden, welche verwendet wird.

<!-- grade_id: cell-4047cf5f2ead4f96, points: 0, grade: false, solution: true -->
```python
import urllib3
import re
import pandas as pd

def grundgesetz() -> dict:
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
print(grundgesetz())
```

Output:
```
File "<ipython-input-62-7c8bebd85a20>", line 10
    print(grundgesetz())
        ^
IndentationError: expected an indented block
```

<!-- grade_id: cell-616e70e5b9f44cb6, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal
import json

def myHash(text: str):
    hash=0
    for ch in text:
        hash = ( hash*281  ^ ord(ch)*997) & 0xFFFFFFFF
    
    return hex(hash)[2:].upper().zfill(8)

assert_equal(myHash(json.dumps(grundgesetz())), '7D877BCB')
```