# Aufgabe 3

Vervollständie die folgende Funktion, welche eine Datei einließt und den Inhalt zeilenweise in Großbuchstaben als Array zurückgibt.

<!-- grade_id: cell-584efe4d9cab9374, points: 0, grade: false, solution: true -->
```python
def get_words(datei: str) -> list:
    liste = []
                
    ### BEGIN SOLUTION
    ### END SOLUTION
        
    return liste
            
print(get_words('testfile.txt'))
```

Output:
```
['AUTO', 'BAUM', 'KATZE', 'HAUS']
```

<!-- grade_id: cell-4e4e63be2c09ea8e, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(hash(frozenset(get_words('german_words.txt'))), 4054237779725047653)
```

```python

```