# Aufgabe 3

Vervollständie die folgende Funktion, welche eine Datei einließt und den Inhalt zeilenweise in Großbuchstaben als Array zurückgibt. Achte darauf, dass leere Zeilen nicht übernommen werden sollen.

<!-- grade_id: cell-584efe4d9cab9374, points: 0, grade: false, solution: true -->
```python
def get_words(datei: str) -> list:
    liste = []
                
    ### BEGIN SOLUTION
    ### END SOLUTION
        
    return liste
            
print(get_words('testfile.txt'))
```

Output:
```
['AUTO', 'BAUM', 'KATZE', 'HAUS']
```

<!-- grade_id: cell-4e4e63be2c09ea8e, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal
import json

def myHash(text: str):
    hash=0
    for ch in text:
        hash = ( hash*281  ^ ord(ch)*997) & 0xFFFFFFFF
    
    return hex(hash)[2:].upper().zfill(8)

assert_equal(myHash(json.dumps(get_words('testfile.txt'))), '633C5C28')
assert_equal(myHash(json.dumps(get_words('german_words.txt'))), '1DF06A3E')
```