# Aufgabe 3

Vervollständie die folgende Funktion, welche eine Datei einließt, in welcher E-Mails gespeichert sind. Die Datei ist in folgendem Format geschrieben:

~~~~
From max.mustermann@fh-swf.de Sat Jan 5 09:14:16 2019
Message
Hallo,

das ist meine super E-Mail.
From susanne.mustermann@fh-swf.de Sun Jan 6 16:31:02 2019
Message
Danke für die E-Mail!
~~~~

Gib ein Tuple zurück, welche als ersten Wert eine Liste aller E-Mail-Adressen beinhaltet und als zweiten Wert die Anzahl aller E-Mail-Adressen, welche aus dem Jahr 2019 stammen.

Beispiel Rückgabe mit der "Datei", welche oben als Beispiel verwendet worden ist: (["max.mustermann@fh-swf.de", "susanne.mustermann@fh-swf.de"], 2)<br>

<!-- grade_id: cell-ea23f96a318a6cd8, points: 0, grade: false, solution: true -->
```python
def get_email_addresses(datei: str) -> tuple:
    email_addresses: list = []
    year_2019: int = 0
    
    ### BEGIN SOLUTION
    ### END SOLUTION
    
    return (email_addresses, year_2019)
    
print(get_email_addresses("email_test.txt"))
```

Output:
```
([], 0)
```

<!-- grade_id: cell-9a32799c5464a003, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal
import json

def myHash(text: str):
    hash=0
    for ch in text:
        hash = ( hash*281  ^ ord(ch)*997) & 0xFFFFFFFF
    
    return hex(hash)[2:].upper().zfill(8)

assert_equal(myHash(json.dumps(get_email_addresses("email_test.txt"))), '601BD29B')
assert_equal(myHash(json.dumps(get_email_addresses("emails.txt"))), '55A13363')
```

Output:
```
/opt/tljh/user/lib/python3.7/unittest/case.py in assertMultiLineEqual(self, first, second, msg)
   1218             diff = '\n' + ''.join(difflib.ndiff(firstlines, secondlines))
   1219             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1220             self.fail(self._formatMessage(msg, standardMsg))
   1221 
   1222     def assertLess(self, a, b, msg=None):

/opt/tljh/user/lib/python3.7/unittest/case.py in fail(self, msg)
    678     def fail(self, msg=None):
    679         """Fail immediately, with the given message."""
--> 680         raise self.failureException(msg)
    681 
    682     def assertFalse(self, expr, msg=None):

AssertionError: '05925D34' != '601BD29B'
- 05925D34
+ 601BD29B
```