# Aufgabe 3

Vervollständie die folgende Funktion, welche eine Datei einließt, in welcher E-Mails gespeichert sind. Die Datei ist in folgendem Format geschrieben:

~~~~
From max.mustermann@fh-swf.de Sat Jan 5 09:14:16 2019
Message
Hallo,

das ist meine super E-Mail.
From susanne.mustermann@fh-swf.de Sun Jan 6 16:31:02 2019
Message
Danke für die E-Mail!
~~~~

Gib ein Tuple zurück, welche als ersten Wert eine Liste aller E-Mail-Adressen beinhaltet und als zweiten Wert die Anzahl aller E-Mail-Adressen, welche aus dem Jahr 2019 stammen.

Beispiel Rückgabe mit der "Datei", welche oben als Beispiel verwendet worden ist: (["max.mustermann@fh-swf.de", "susanne.mustermann@fh-swf.de"], 2)<br>

<b>Hinweis:</b> Zur automatischen Korrektur, muss die Liste als Frozenset innerhalb des Tupels zurückgegeben werden. Die Liste "email_addresses" kann innerhalb der Funktion jedoch normal verwendet werden.

<!-- grade_id: cell-ea23f96a318a6cd8, points: 0, grade: false, solution: true -->
```python
def get_email_addresses(datei: str) -> tuple:
    email_addresses: list = []
    year_2019: int = 0
    
    ### BEGIN SOLUTION
    ### END SOLUTION
                
    
    return (frozenset(email_addresses), year_2019)
    
print(get_email_addresses("email_test.txt"))
```

Output:
```
(frozenset(), 0)
```

<!-- grade_id: cell-9a32799c5464a003, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(hash(get_email_addresses("email_test.txt")), -7979767667009334678)
assert_equal(hash(get_email_addresses("emails.txt")), -3397707965979964180)
```